Sleep need driven oscillation of glutamate synaptic phenotype

Author response:

The following is the authors’ response to the previous reviews.

eLife Assessment

This important study showing that sleep deprivation increases functional synapses while depleting silent synapses supports previous findings that excitatory signaling increases during wakefulness. This manuscript focuses in particular on AMPA/NMDA ratios. An interesting, although speculative, aspect of the manuscript is the inclusion of a model for the accumulation of sleep need that is based upon the MEF2C transcription factor but also links to the sleep-regulating SIK3-HDAC4/5 pathway. The authors have clarified some questions raised in the previous review, but the evidence for major claims was still found to be incomplete, requiring additional experimentation.

The major claims of this study are: 1) SD increases the AMPA/NMDA receptor ratio and RS restores it; 2) SD decreases silent synapses compared to CS and RS restores their number after SD; 3) the majority of SD-induced DEGs are found in ExIT cells (glutamate pyramidal neurons projecting within the telencephalon); 4) ExIT SD-induced DEGs are enriched for genes encoding synaptic shaping components and for autism spectrum disorder risk and; 5) these DEGs are also enriched for DEGs induced by Mef2c loss of function restricted to forebrain glutamate neurons (ExIT cells comprise a subset of these) and by over-expression of constitutively nuclear HDAC4 that represses MEF2c transcriptional function. The last claim is consistent with an intracellular signaling model (presented as a hypothesis to be tested, in figure 4B).

[The above is added to the start of the discussion section.]

The specific claims are supported by solid evidence provided in this manuscript. The statistical support is now more clearly presented, with several changes in response to queries by reviewer 1.

The technical issues raised by reviewer 1 do not detract from the claims, thus supported. The rationale for this assessment is expanded below in response to reviewer 1.

Summary:

This manuscript by Vogt et al examines how the synaptic composition of AMPA and NMDA receptors changes over sleep and wake states. The authors perform whole-cell patch clamp recordings to quantify changes in silent synapse number across conditions of spontaneous sleep, sleep deprivation, and recovery sleep after deprivation. They also perform single nucleus RNAseq to identify transcriptional changes related to AMPA/NMDA receptor composition following spontaneous sleep and sleep deprivation. The findings of this study are consistent with a decrease in silent synapse number during wakefulness and an increase during sleep. However, these changes cannot be conclusively linked to sleep/wake states. Measurements were performed in motor cortex, and sleep deprivation was achieved by forced locomotion, raising the possibility that recent patterns of neuronal activity, rather than sleep/wake states, are responsible for the observed results.

Strengths:

This study examines an important question. Glutamatergic synaptic transmission has been a focus of studies in the sleep field, but AMPA receptor function has been the primary target of these studies. Silent synapses, which contain NMDA receptors but lack AMPA receptors, have important functional consequences for the brain. Exploring the role of sleep in regulating silent synapse number is important to understanding the role of sleep in brain function. The electrophysiological approach of measuring the failure rate ratio, supported by AMPA/NMDA ratio measurements, is a rigorous tool to evaluate silent synapse number.

The authors also perform snRNAseq to identify genes differentially expressed in the spontaneous sleep and sleep deprivation groups. This analysis reveals an intriguing pattern of upregulated genes controlled by HDAC4 and Mef2c, along with synaptic shaping component genes and genes associated with autism spectrum disorder, across cell types in the sleep deprivation group. This unbiased approach identifies candidate genes for follow-up studies. The finding that ASD-risk genes are differentially expressed during SD also raises the intriguing possibility that normal sleep function is disrupted in ASD.

Weaknesses:

A major consideration to the interpretation of this study is the use of forced locomotion for sleep deprivation. Measurements are made from motor cortex, and therefore the effects observed could be due to differences in motor activity patterns across groups, rather than lack of sleep per se.

Experimentally induced lack of sleep always involves differences in motor activity. As previously noted in revision 1, motor learning is unlikely to occur in this paradigm and inspection of the video (in supplementary materials) shows no repetitive motor behavioral sequences during the sleep deprivation, nor can this be considered exercise due to the very slow speed of treadmill movement employed. The obvious major difference between groups is a lack of sleep per se. (See below in the “Recommendations for authors”, reviewer 1 for comments on localized wake activity inducing localized sleep-need responses)

Considering that other groups have failed to find a difference in AMPA/NMDA ratio in mice with different spontaneous sleep/wake histories (Bridi et al., Neuron 2020), confirmation of these findings in a different brain region would greatly strengthen the study.

The study of Bridi et al., Neuron 2020, is not comparable to our study for several important reasons. First, their compared groups were from different circadian phases (180 degrees out of phase), whereas in our study, the circadian times for each group were matched (ZT=6hours). Second, experimentally induced sleep loss did not occur whereas it was a focus of our study. Third, spontaneous sleep/wake cannot be accurately matched amongst subjects whereas in our study, sleep loss was matched exactly between groups.

We agree that assessment of AMPA/NMDA ratio and silent synapse number in sleep deprived compared to ad libitum sleep in other areas of the neocortex is of great interest and something we hope to pursue. It would not be surprising to find differences as preliminarily reported by Bahl, et al., Nat Commun. 2024 Jan 26;15(1):779. However, such data would not further strengthen our already well supported evidence for the differences we report in the motor cortex.

The electrophysiological measurements and statistical analyses raise several questions. Input resistance (cutoffs and actual values) are not provided, making it difficult to assess recording quality.

As stated in our first reply, these data were omitted (an admitted oversight on our part) but are now supplied in the methods section as, “Series resistance values for the recording pipette ranged between 8 and 15 MOhm and experiments with changes larger than 25% were not used for further analyses”. We have now also added the Rs/Rm (as a separate column) for each recorded neuron in table 1.

Parametric one-way ANOVAs were used, although the data do not appear to be normally distributed.

We have now removed all the One-way ANOVA tests for clarity (non-parametric tests were previously supplied in addition to the one-way ANOVA tests). Determination of significance with Kruskal-Wallis non-parametric test has not altered statistical support for our conclusions.

Reviewer 1 correctly points out that we had not tested for normality of our distributions- the distributions are likely to be normal but the sample size is too small to confidently make this call for the ratio data which is why we removed the one-way ANOVA’s entirely from table 1.

Two-way ANOVA’s are used to assess AMPA and EPSC amplitudes and failure rates (table 1 tab 2&5) across sleep conditions. As now indicated (table 1, tab 2&5), the distributions of AMPA and NMDA amplitudes and FRs passed the D'Agostino & Pearson test for normality and QQ plots provide illustration supporting this claim.

In addition, for the AMPA/NMDA and FRR measurements (Figures 1E, F), the SD group (rather than the control sleep group) was used as the control group for post-hoc comparisons, but it is unclear why.

The label of “control group” is arbitrary. CS and RS groups are similar (sleep density for RS>CS as expected). Since this appears to be confusing, we now compare all groups to one another in table 1 with the same statistical outcome (additional comparison of CS to RS).

While the data appear in line with the authors' conclusions, the number of mice (3/group) and cells recorded is low, and adding more would better account for inter-animal variability and increase the robustness of the findings.

Of course, the larger the sample, the better the approximation to the population. Our sample sizes yielded significant differences at the usual p<=0.05 threshold with non-parametric testing. A larger sample size could allow for normality testing of the distributions of the data, but fortunately, this was not necessary to support our conclusions.

The snRNAseq data are intriguing. However, several genes relevant to the AMPA/NMDA ratio are mentioned, but the encoded proteins would be expected to have variable effects on AMPA/NMDA receptor trafficking and function, making the model presented in Figure 4C oversimplified. A more thorough discussion of the candidate genes and pathways that are upregulated during sleep deprivation, the spatiotemporal/posttranslational control of protein expression, and their effects on AMPA/NMDA trafficking vs function is warranted.

We have not studied the candidate genes at this point and do not yet understand their potential role(s) in sleep-related AMPA/NMDA functional ratio, only that their expression levels are altered with sleep condition. We agree with the reviewer that the data are intriguing and in need of further investigation. An important first step that can help direct such studies is the identification and preliminary characterization of good candidate genes with respect their cell type specificity, significance and fold change as we have done. Their potential roles likely depend on “the spatiotemporal/posttranslational control” and other factors as reviewer 1 notes.

Reviewer #2 (Public review):

Here Vogt et al., provide new insights into the need for sleep and the molecular and physiological response to sleep loss. The authors expand on their previously published work (Bjorness et al., 2020) and draw from recent advances in the field to propose a neuron-centric molecular model for the accumulation and resolution of sleep need and basis of restorative sleep function. While speculative, the proposed model successfully links important observations in the field and provides a framework to stimulate further research and advances on the molecular basis of sleep function. In my review, I highlight the important advances of this current work, the clear merits of the proposed model, and indicate areas of the model that can serve to stimulate further investigation.

Strengths:

Reviewer comment on new data in Vogt et al., 2024

Using classic slice electrophysiology, the authors conclude that wakefulness (sleep deprivation (SD)) drives a potentiation of excitatory glutamate synapses, mediated in large part by "un-silencing" of NMDAR-active synapses to AMPAR-active synapses. Using a modern single nuclear RNAseq approach the authors conclude that SD drives changes in gene expression primarily occurring in glutamatergic neurons. The two experiments combined highlight the accumulation and resolution of sleep need centered on the strength of excitatory synapses onto excitatory neurons. This view is entirely consistent with a large body of extant and emerging literature and provides important direction for future research.

Consistent with prior work, wakefulness/SD drives an LTP-type potentiation of excitatory synaptic strength on principle cortical neurons. It has been proposed that LTP associated with wake, leads to the accumulation of sleep need by increasing neuronal excitability, and by the "saturation" of LTP capacity. This saturation subsequently impairs the capacity for further ongoing learning. This new data provides a satisfying mechanism of this saturation phenomenon by introducing the concept of silent synapses. The new data show that in mice well rested, a substantial number of synapses are "silent", containing an NMDAR component but not AMPARs. Silent synapses provide a type of reservoir for learning in that activity can drive the un-silencing, increasing the number of functional synapses. SD depletes this reservoir of silent synapses to essentially zero, explaining how SD can exhaust learning capacity. Recovery sleep led to restoration of silent synapses, explaining how recovery sleep can renew learning capacity. In their prior work (Bjorness et al., 2020) this group showed that SD drives an increase in mEPSC frequency onto these same cortical neurons, but without a clear change in pre-synaptic release probability, implying a change in the number of functional synapses. This prediction is now born out in this new dataset.

The new snRNAseq dataset indicates the sleep need is primarily seen (at the transcriptional level) in excitatory neurons, consistent with a number of other studies. First, this conclusion is corroborated by an independent, contemporary snRNAseq analysis recently available as a pre-print (Ford et al., 2023 BioRxiv https://doi.org/10.1101/2023.11.28.569011). A recently published analysis on the effects of SD in drosophila imaged synapses in every brain region in a cell-type dependent manner (Weiss et al., PNAS 2024), concluding that SD drives brain wide increases in synaptic strength almost exclusively in excitatory neurons. Further, Kim et al., Nature 2022, heavily cited in this work, show that the newly described SIK3-HDAC4/5 pathway promotes sleep depth via excitatory neurons and not inhibitory neurons.

The new experiments provided in Fig1-3 are expertly conducted and presented. This reviewer has no comments of concern regarding the execution and conclusions of these experiments.

Reviewer comment on model in Vogt et al., 2024

To the view of this reviewer the new model proposed by Vogt et al., is an important contribution. The model is not definitively supported by new data, and in this regard should be viewed as a perspective, providing mechanistic links between recent molecular advances, while still leaving areas that need to be addressed in future work. New snRNAseq analysis indicates SD drives expression of synaptic shaping components (SSCs) consistent with the excitatory synapse as a major target for the restorative basis of sleep function. SD induced gene expression is also enriched for autism spectrum disorder (ASD) risk genes. As pointed out by the authors, sleep problems are commonly reported in ASD, but the emphasis has been on sleep amount. This new analysis highlights the need to understand the impact on sleep's functional output (synapses) to fully understand the role of sleep problems in ASD.

Importantly, SD induced gene expression in excitatory neurons overlap with genes regulated by the transcription factor MEF2C and HDAC4/5 (Fig. 4). In their prior work, the authors show loss of MEF2C in excitatory neurons abolished the SD transcriptional response and the functional recovery of synapses from SD by recovery sleep. Recent advances identified HDAC4/5 as major regulators of sleep depth and duration (in excitatory neurons) downstream of the recently identified sleep promoting kinase SIK3. In Zhou et al., and Kim et al., Nature 2022, both groups propose a model whereby "sleep-need" signals from the synapse activate SIK3, which phosphorylates HDAC4/5, driving cytoplasmic targeting, allowing for the de-repression and transcriptional activation of "sleep genes". Prior work shows that HDAC4/5 are repressors of MEF2C. Therefore, the "sleep genes" derepressed by HDAC4/5 may be the same genes activated in response to SD by MEF2C. The new model thereby extends the signaling of sleep need at synapses (through SIK3-HDAC4/5) to the functional output of synaptic recovery by expression of synaptic/sleep genes by MEF2C. The model thereby links aspects of expression of sleep need with the resolution of sleep need by mediating sleep function: synapse renormalization.

Weaknesses:

Areas for further investigation.

In the discussion section Vogt et al., explore the links between excitatory synapse strength, arguably the major target of "sleep function", and NREM slow-wave activity (SWA), the most established marker of sleep need. SIK3-HDAC4/5 have major effects on the "depth" of sleep by regulating NREM-SWA. The effects of MEF2C loss of function on NREM SWA activity are less obvious, but clearly impact the recovery of glutamatergic synapses from SD. The authors point out how adenosine signaling is well established as a mediator of SWA, but the links with adenosine and glutamatergic strength are far from clear. The mechanistic links between SIK3/HDAC4/5, adenosine signaling, and MEF2C, are far from understood. Therefore, the molecular/mechanistic links between a synaptic basis of sleep need and resolution with NREM-SWA activity require further investigation.

Additional work is also needed to understand the mechanistic links between SIK3-HDAC4/5 signaling and MEF2C activity. The authors point out that constitutively nuclear (cn) HDAC4/5 (acting as a repressor) will mimic MEF2C loss of function. This is reasonable, however, there are notable differences in the reported phenotypes of each. Notably, cnHDAC4/5 suppresses NREM amount and NREM SWA but had no effect on the NREM-SWA increase following SD (Zhou et al., Nature 2022).

We speculate that the effect of cnHDAC4/5 to reduce NREM-SWA together with the reduction of NREM amount may be due to a localized increase in neuronal excitability of arousal centers, which would be expected to mask NREM-SWA. Rebound NREM-SWA may reflect the relative rebound increase of NREM-SWA still present under chronic masking conditions (induced by cnHDAC4/5) of increased arousal system excitability. A similar effect to overcome NREM-SWA masking was reported in a Kcna2 KO mouse (a Shaker homologue) by Douglas, et al. (2007, BMC Biol).

Loss of MEF2C in CaMKII neurons had no effect on NREM amount and suppressed the increase in NREM-SWA following SD (Bjorness et al., 2020). These instances indicate that cnHDAC4/5 and loss of MEF2C do not exactly match suggesting additional factors are relevant in these phenotypes. Likely HDAC4/5 have functionally important interactions with other transcription factors, and likewise for MEF2C, suggesting areas for future analysis.

This is not a surprising outcome since both MEF2c and HDAC4/5 are transcription factors whose function(s) are determined by multiple other factors a subset of which are relevant to sleep conditions while other determining factors are not necessarily relevant to sleep. These factors can include their phosphorylation state, genomic accessibility, and interaction with other transcription factors. All these other factors are known to be both cell type specific and determined by intracellular conditions, that in turn, are affected by extracellular conditions and ligands. We certainly agree there is much future analysis needed.

One emerging theme may be that the SIK3-HDAC4/5 axis are major regulators of the sleep state, perhaps stabilizing the NREM state once the transition from wakefulness occurs. MEF2C is less involved in regulating sleep per se, and more involved in executing sleep function, by promoting restorative synaptic modifications to resolve sleep need.

A useful way to restate the above might be to distinguish between control of arousal levels determining the behavioral states, wake or sleep (including REM sleep) and control of sleep function. The term, sleep, is typically used to describe the behavioral state of sleep that acts as a permissive gate to sleep function (that resolves sleep need). The sleep state should not be conflated with sleep function. There is abundant evidence that control of arousal can be dissociated from sleep need and sleep function.

Finally, advances in the roles of the respective SIK3-HDAC4/5 and MEF2C pathways point towards transcription of "sleep genes", as clearly indicated in the model of Fig.4. Clearly more work is needed to understand how the expression of such genes ultimately lead to resolution of sleep need by functional changes at synapses.

We are in full agreement. We also note the SIK3-HDAC4/5 pathway may have more than one role, i.e., to affect arousal centers to alter behavioral state and, more generally, to control MEF2c’s transcriptional activity thus controlling sleep-related, glutamate, synaptic phenotype.

What are these sleep genes and how do they mechanistically resolve sleep need? Thus, the current work provides a mechanistic framework to stimulate further advances in understanding the molecular basis for sleep need and the restorative basis of sleep function.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Major comments:

(1) I appreciate the authors' thoughtful discussion of the use of forced locomotion for their sleep deprivation technique in their response, as well as the additional information that was provided regarding use of the treadmill in the manuscript. However, given that previous studies have failed to find a difference in AMPA/NMDA ratio following spontaneous sleep vs wake, confirmation of the findings in a non-motor brain region with the same SD technique (or confirmation within motor cortex with a different technique, although the authors correctly point out that other techniques also increase locomotor activity) would greatly strengthen the paper.

Addressed above

Notably, differences in motor activity patterns, not necessarily overall amount of locomotion, may induce differential synaptic changes between groups. This point at least warrants acknowledgement and discussion, but this has not been incorporated into the text of the manuscript.

We will incorporate the following into the discussion:

There is evidence that learning of a motor task or experience of forced altered motor activity can result in localized increases in NREM (slow wave sleep)-slow wave activity (Huber R, Ghilardi MF, Massimini M, Tononi G. Local sleep and learning. Nature. 2004;430(6995):78-81); Huber et al., 2006) in the motor cortex. Since SWS-SWA is considered a marker for sleep homeostasis, the altered motor activity induced increase of SWS-SWA was considered evidence for sleep-related function. Our earlier work has clearly shown that the treadmill method of SD increases frontal cortical SWS-SWA rebound, indicating a sleep-homeostatic process (Bjorness et al., 2016; Bjorness et al., 2020). Furthermore, we have also shown that this means of experimental SD causes similar glutamate synaptic changes as those observed using other means of SD like gentle handling (Liu, et al., JoNS 2010).

(2) The number of mice and cells used for electrophysiology in this study remains low; more animals should be included to account for inter-animal variability.

For this study, increasing the number of mice and cells will have p<0.05 chance of altering our conclusions by rejecting the null hypotheses of the electrophysiology findings.

(3) The additional methodological information provided allays some of my concerns regarding the electrophysiological data. However, information about the input resistance (cutoffs used and/or actual values) is still not provided, which is important for assessing recording quality.

We have now supplied the experimentally determined input resistance for each neuron used in this study (a separate column in table 1, tabs marked, “data”).

(4) It is not meaningful to compare raw AMPA or NMDA responses because stimulus electrode placement will differ between cells, potentially activating different numbers of afferents. Presenting these comparisons (Figure 1C) has the potential to mislead the reader.

This is not misleading (it didn’t mislead reviewer 1) as we described the conditions. As expected by reviewer 1, the variability using “raw AMPA or NMDA responses…” was too great, but did indicate an interaction between receptor responses and sleep condition. This provided (as stated in the results section) rationale to examine, and to only draw conclusions from the AMPA/NMDA amplitude and FR ratios.

(5) I appreciate clarification on the statistics and the authors' response has answered some of my questions. However, this also raises additional questions. What test was used to determine normality (and therefore whether to perform a parametric vs nonparametrictest)?

Described above.

Why was the FRR data analysis changed to a parametric test, when it does not appear that the data are normally distributed?

Showing the parametric test was a mistake on our part- there are not enough samples to conclusively conclude the distributions are normal as reviewer 1 correctly suspects. However, the non-parametric Kruskal-Wallis tests that we also show in table 1 indicate significant differences between conditions and the non-parametric, two-stage linear step-up procedure of Benjamini, Krieger and Yekutieli, indicates significant differences between CS-SD and RS-SD but not for CS-RS, supporting our conclusions. The (unsupported) parametric tests are now removed in Table 1 leaving behind the non-parametric test.

Why were post-hoc tests chosen to compare to a control group rather than all pairwise comparisons,

We now provide post-hoc all-pairwise comparisons to give the same results using the BKY analysis.

and why was the SD rather than CS group used as the control in Figures 1E and F?

Why were different post-hoc tests chosen for the data in Figures 1E, F?

There was no need for this and we now, only show statistics that are used to draw our conclusions for the AMPA/NMDA EPSC ratios data shown in Figure 1E and Failure Rate Ratios data shown in Figure 1F (the conclusions are supported by the non-parametric post-hoc test and remain unchanged).

(6) Genes in the SSC, ASD, Mef2cKO, and HD4cn categories are almost exclusively upregulated in the SD group compared to the CS group (Figure 4A). As the authors point out in their response, "No claim of mechanism linking the changed expression to altered AMPAR or NMDAR activity can be made at this point," largely due to the fact that we do not know the spatiotemporal or posttranslational modification patterns of the translated proteins, and how they affect receptor trafficking vs function. This is in agreement with my original point: as written (and as illustrated in Figure 4C), the manuscript implies that upregulation during SD increases the AMPA/NMDA ratio via receptor trafficking,

The model indicates a likely (but not necessarily exclusive) role for AMPA/NMDA trafficking to explain the functional electrophysiological data that we do report and which is not in dispute. The SSC-DEGs in ExIT cells are consistent with sleep-altered AMPA/NMDA trafficking but remain only a correlation. However, the point is taken and Figure 4c has been revised to only reflect what we have observed electrophysiologically and the speculated mechanism(s) mediated by observed SSC-DEGs are illustrated with “?’s”.

while in reality the picture is likely much more complicated, and therefore a more thorough discussion is warranted. Some discussion was provided in the authors' response but does not appear to have been incorporated into the text or Figure 4C.

As indicated above the proposed model is changed in Figure 4c to more explicitly indicate which aspects reflect our electrophysiological data and which aspects reflect only an association of observations.

Minor comments:

(1) Please justify only using male mice

We had to start somewhere with our limited resources. Our intentions are to follow up with similar experiments using female mice, should funding be realized.

(2) The model in Figure 4C is oversimplified and remains problematic, for the reasons stated in comment #6, above.

See responses above.

(3) Figure 4D remains confusing

We agree. The unnecessary addition of adenosine effects on cholinergic arousal centers (experimentally well supported), have been removed from the figure to provide a more focused indication of how SWS-SWA can be related to either MEF2c and/or to ADORA1 activation through reduction of glutamate synaptic strength. ADORA1 activation elicits reduced glutamate synaptic activity through pre- and postsynaptic inhibition whereas MEF2c activation is essential to reduce sleep elicited, glutamate EPSC reduction. Reduced glutamate synaptic strength, whatever the cause, is associated with increased SWS-SWA.

Changes from the earlier version to the revised version

Abstract

Abstract

Sleep loss increases AMPA-synaptic strength and number in the neocortex. However, this is only part of the synaptic sleep loss response. We report increased AMPA/NMDA EPSC ratio in frontal-cortical pyramidal neurons of layers 2-3. Silent synapses are absent, decreasing the plastic potential to convert silent NMDA to active AMPA synapses. These sleep loss changes are recovered by sleep. Sleep genes are enriched for synaptic shaping cellular components controlling glutamate synapse phenotype, overlap with autism risk genes and are primarily observed in excitatory pyramidal neurons projecting intra-telencephalically. These genes are enriched with genes controlled by the transcription factor, MEF2c and its repressor, HDAC4. Sleep genes can thus provide a framework within which motor learning and training occurs mediated by sleep-dependent oscillation of glutamate-synaptic phenotypes.

Added in the revised version: Introduction

Removed in the revised version: Main Text

CNS control of arousal maintains waking neuronal activities needed for foraging, danger avoidance and reproductive behaviors. Reduction of arousal to unconsciousness can gate sleep, however, the function served by this sleep, in contrast to wake function, remains enigmatic. A reasonable teleologically oriented goal is more effective future foraging, danger avoidance and reproductive behaviors, based on procedural and episodic learning from past waking experience. Effective learning involves memory of waking experience and underlying long-term plasticity (LTP) (Moser et al., 1998; Morris et al., 2003; Stickgold, 2005). At the cellular level, extensive activity dependent LTP of cortical glutamate synapses can lead to a well-documented increase of synaptic strength of cortical glutamatergic synapses during waking. It is reflected by increased frequency and amplitude of glutamatergic miniature synaptic currents (mEPSC) that is restored by subsequent sleep (Liu et al., 2010; Torrado Pacheco et al., 2021), consistent with the synaptic homeostasis hypothesis of sleep (SHY) (Tononi and Cirelli, 2014).

Our recent findings indicate that during waking, not only are individual synapses strengthened (overall glutamate mEPSC amplitude in layer 2-3 pyramidal cells is increased by SD) but, also, the number of electro-physiologically active synapses is increased (increased frequency in correlation with unaltered or decreased probability of release) (Bjorness et al., 2020). The molecular underpinnings of these changes in glutamate synaptic strength and number were examined to determine the changes in the glutamate synaptic phenotype and what the implication(s) of these changes, if any, might be with regards to propensity for synaptic plasticity (meta-plasticity).

We investigated sleep/wake mediated electrophysiological changes in glutamatergic synapses in motor cortex, layer 2-3, pyramidal neurons since our previous work demonstrates wake/sleep up-regulation and down-regulation, respectively, in these cells’ glutamate synapses (Bjorness et al., 2020). We observed a wake/sleep oscillation of AMPA/NMDA synaptic response ratio in pyramidal neurons of motor neocortex coupled with an oscillating fraction of silent synapses (NMDAR-mediated synaptic response but no AMPAR mediated response) in response to sleep need. We then, extended these physiological findings to include characterization of sleep-need related differential expression of glutamatergic synaptic shaping genes that can control the observed oscillating glutamate synapse phenotype.

Results

Sleep need driven, glutamate, synaptic phenotype

Earlier version: Three cohorts of C57 BL/6 male mice were allowed either to 1) sleep ad lib (control sleep, CS); 2) prevented from sleeping from ZT=0 to 6h (sleep deprived, SD); or, 3) after 4h SD allowed recovery sleep, ad lib for 2h (RS). Ex Vivo acute brain slices from the motor cortex were then prepared from these animals at ZT=6h and whole cell recordings obtained from layer 2/3 pyramidal neurons (see methods for details). Locally evoked glutamatergic EPSCs recorded at Vhold = -90mV had little NMDAR contribution due to Mg++ block. These were compared in the same cell, to EPSCs evoked at Vhold = +50mV with little AMPAR contribution due to reduced AMPAR driving force, but, having large NMDAR currents due to voltage sensitive relief from Mg++ blockade (NMDAR currents had reversed polarity since ENMDAR =∼0mV). For each recorded neuron (one neuron/ slice/ animal) an average response to repeated stimuli (>10 stimuli, 100pA, applied every 10sec) to elicit a EPSC was determined at each holding potential. A two-way ANOVA analysis showed a significant interaction between AMPA matched to NMDA EPSC response for each neuron, and sleep condition (F (2, 21) = 7.268, p<0.004; Figure 1 A, C, E). When considered independently, neither the effect of sleep condition nor of EPSC subtype reached significance at p<0.05 (Figure 1 C).

Revised version: Three cohorts of C57 BL/6 male mice were allowed either to 1) sleep ad lib (control sleep, CS); 2) prevented from sleeping from ZT=0 to 6h (sleep deprived, SD); or, 3) after 4h SD allowed recovery sleep, ad lib for 2h (RS). Ex Vivo acute brain slices from the motor cortex were then prepared from these animals at ZT=6h and whole cell recordings obtained from layer 2/3 pyramidal neurons (see methods for details). Locally evoked glutamatergic EPSCs recorded at Vhold = - 90mV had little NMDAR contribution due to Mg++ block. These were compared in the same cell, to EPSCs evoked at Vhold = +50mV with little AMPAR contribution due to reduced AMPAR driving force, but, having large NMDAR currents due to voltage sensitive relief from Mg++ blockade (NMDAR currents had reversed polarity since ENMDAR =∼0mV). For each recorded neuron (one neuron/ slice/ animal) an average response to repeated stimuli (>10 stimuli, 100pA, applied every 10sec) to elicit a EPSC was determined at each holding potential. A two-way ANOVA analysis showed a significant interaction between AMPA matched to NMDA EPSC response for each neuron, and sleep condition (F (2, 21) = 7.268, p<0.004; Figure 1 A, C, E; Table 1, tabs1&2). When considered independently, neither the effect of sleep condition nor of EPSC subtype reached significance at p<0.05 Figure 1 C). To control for neuron-to-neuron EPSC variability (due to variability of stimulation-evoked afferent activation), we calculated the AMPA/NMDA ratio for three sleep conditions, CS, SD and RS. We observed a significant difference between sleep condition cohorts non-parametric Kruskal-Wallis test, p<0.0022). Significance tests for multiple comparisons between sleep conditions, gave CS-SD, p<0.031; SD-RS, p<0.005; and CS- RS, p<0.191 (Two-stage linear step-up procedure of Benjamini, Krieger and Yekutieli; Table 1, tabs1&3; Figure 1 E).

Removed in the revised version: Sleep need-dependent responses of AMPA/NMDA ratio and silent synapses in motor cortex. A) Examples of AMPA currents at -90 mV holding potential and NMDA currents at +50 mV holding potential are shown for: control sleep (CS), 6 h of sleep deprivation (SD) and 4 h of sleep deprivation followed by 2 h of recovery sleep (RS). Traces (100msec duration) are scaled to the NMDA current for comparison (NMDA current measured @ 40msec after AMPA peak current). B) Examples of successes (black) and failures (red) at -90 mV (left, AMPA) and +50 mV (right, NMDA) after minimal stimulation of excitatory inputs to motor cortex pyramidal neurons are shown; Top row: CS sleep, middle row: SD and bottom row: RS (Cal. 10pAX20msec). Rate of failures (% of all stimuli delivered) for AMPA EPSCs (blue) and NMDA EPSCs (red) in the three conditions (+/- sem). C), D) Average (+/- sem) AMPAR and NMDAR EPSC responses (unmatched) and failure rates, respectively, for each sleep condition. E), F) Matched AMPA/NMDA EPSC response and AMPA/NMDA failure rate ratios, respectively, are shown for the three conditions (N=1 cell/slice, 2-3 slice/animal, 3 animals/condition).

Removed in the revised version: To control for neuron-to-neuron EPSC variability (due to variability of stimulation-evoked afferent activation), we calculated the AMPA/NMDA ratio, matched for each recorded neuron. This showed a significantly greater AMPA/NMDA EPSC ratio (F (2, 21)= 3.203) and comparing SD to CS or RS, p<=0.03 or p<=0.02, respectively; Dunnett’s multiple comparisons test; Table S1, tab 2).

Earlier version: Many glutamatergic synapses on cortical pyramidal cells do not respond to activation at resting membrane potential, yet they do have an NMDAR synaptic component that can be observed at Vhold=+50mV. After activation by LTP-evoking stimulation, these ‘silent’ synapses can convert to active, non-silent synapses by virtue of LTP-induced AMPAR insertion into the post-synaptic active zone (Liao et al., 1995). The ratio of AMPAR component failure rate (Vhold= -90mv; FR-90 for active synapses, As) matched with the same neuron’s FR+50 (Vhold= +50mV; for As + silent synapses, Ss) has been used to estimate the change in the fraction of Ss/As in response to LTP (Liao et al., 1995). We employed a similar analysis to examine the effects of SD and sleep on functional synaptic AMPAR and NMDAR composition. We observed a significant interaction between FR’s and sleep conditions, assessed by a two-way ANOVA mixed effects model (F (2, 19) = 5.550, p = 0.0126; Figure 1 B1,2,3 D F; Table S1, tab 4). Multiple comparisons revealed no significant difference between FR-90 and FR+50 in SD indicating far fewer numbers of silent synapses (Ss) in this condition. However, both CS and RS restored the relatively decreased number of Ss compared to SD. FR-90 is significantly smaller than FR+50 for both CS and RS (p<0.0001 and <0.0005, respectively).

Revised version: Many glutamatergic synapses on cortical pyramidal cells do not respond to activation at resting membrane potential, yet they do have an NMDAR synaptic component that can be observed at Vhold=+50mV. After activation by LTP-evoking stimulation, these ‘silent’ synapses can convert to active, non-silent synapses by virtue of LTP-induced AMPAR insertion into the post-synaptic active zone (Liao et al., 1995). The ratio of AMPAR component failure rate (Vhold= -90mv; FR-90 for active synapses, As) matched with the same neuron’s FR+50 (Vhold= +50mV; for As + silent synapses, Ss) has been used to estimate the change in the fraction of Ss/As in response to LTP (Liao et al., 1995). We employed a similar analysis to examine the effects of SD and sleep on functional synaptic AMPAR and NMDAR composition. As with the EPSC amplitude analysis, there was a significant interaction for failure rates at Vhold= -90mv and Vhold= +50mV, and the three sleep conditions (p<0.031; F (1.742, 8.708) = 5.550; Table 1, Tabs 4&5; Figure 1 D). Multiple comparisons (Šídák’s multiple comparisons test; Table 1, Tab 5) of mean FR at Vhold= -90mv (for As) and Vhold= +50mV (for As+Ss) across sleep conditions, revealed no significant difference between FR-90 and FR+50 in SD (p<0.916), indicating far fewer numbers of silent synapses (Ss) in this condition. However, both CS and RS restored the number of Ss. FR-90 is significantly smaller than FR+50 for both CS and RS (p<0.030 and <0.0035, respectively).

Added in the revised version: Sleep need-dependent responses of AMPA/NMDA ratio and silent synapses in motor cortex.

Added in the revised version: A) Examples of AMPA currents at -90 mV holding potential and NMDA currents at +50 mV holding potential are shown for: control sleep (CS), 6 h of sleep deprivation (SD) and 4 h of sleep deprivation followed by 2 h of recovery sleep (RS). Traces (100msec duration) are scaled to the NMDA current for comparison (NMDA current measured @ 40msec after AMPA peak current). B) Examples of successes (black) and failures (red) at -90 mV (left, AMPA) and +50 mV (right, NMDA) after minimal stimulation of excitatory inputs to motor cortex pyramidal neurons are shown; Top row: CS sleep, middle row: SD and bottom row: RS (Cal. 10pAX20msec). Rate of failures (% of all stimuli delivered) for AMPA EPSCs (blue) and NMDA EPSCs (red) in the three conditions (+/-sem). C), D) Average (+/-sem) AMPAR and NMDAR EPSC responses (unmatched) and failure rates, respectively, for each sleep condition. E), F) Matched AMPA/NMDA EPSC response and AMPA/NMDA failure rate ratios, respectively, are shown for the three conditions (N=1 cell/slice, 2-3 slice/animal, 3 animals/condition).

Earlier version: The condition specific ratio of FR+50mV to FR-90mV (FRR) is a power function of the probability of failure for each of the stimulated synapses (1-prelease)stimulated-synapse and thus, inversely proportional (since (1-prelease)<1) to the number of silent synapses (Ss). A binomial probability distribution model is employed with the assumptions that: 1) each stimulus activates the same set of pre-synaptic terminals for a given neuron; 2) each terminal has a constant prelease; 3) P(failure) for a given set of stimulus-activated terminals = (1-prelease)#activated terminals. The FR+50mV/FR-90mV = (1-prelease)As+Ss/(1-prelease)As and ln(FR+50mV)/ln(FR-90mV) =(As+Ss)/As = 1+Ss/As. SD’s FR+50mV/FR-90mV=∼1, indicating ∼no silent synapses. The FRR of Ss/As for CS=1.44 and for RS=1.46, compared to SD=0.09. The difference between sleep conditions and FRR is significant (F (2, 19) = 11.3, Table S1, tab5). Multiple comparisons (Holm-Sidak, Table S1, tab5) indicate the near absence of silent synapses was reversed by either CS or RS (SDxCS; p<0.0011 and SDxRS: p<0.0006; Table S1, tab 5; Figure 1 F). This suggests there are few if any silent synapses available for conversion by LTP to active synapses following SD and this availability is recovered by sleep.

Revised version: The condition specific ratio of FR+50mV to FR-90mV (FRR) is a ratio of power functions of the probability of failure for each of the stimulated synapses (1-prelease)stimulated-synapse and thus, inversely proportional (since (1-prelease)<1) to the number of silent synapses (Ss). A binomial probability distribution model is employed with the assumptions that: 1) each stimulus activates the same set of pre-synaptic terminals for a given neuron; 2) each terminal has a constant prelease; 3) P(failure) for a given set of stimulus-activated terminals = (1-prelease)#activated terminals. The FRR for FR+50mV/FR-90mV = (1-prelease)As+Ss/(1-prelease)As and ln(FR+50mV)/ln(FR-90mV) =(As+Ss)/As=1+Ss/As. SD’s FR+50mV/FR-90mV=∼1, indicating ∼no silent synapses. Ss/As for CS=1.44 and for RS=1.46, compared to SD=0.09. As with the AMPA/NMDA RR, we observed a significant difference for FRR between sleep condition cohorts (non-parametric Kruskal-Wallis test, p<0.0006). Significance tests for multiple comparisons between sleep conditions, gave CS-SD, p<0.02; SD-RS, p<0.0012; and CS-RS, p<0.366 (Two-stage linear step-up procedure of Benjamini, Krieger and Yekutieli; Table 1, tabs 4&6; Figure 1 F) This suggests there are few if any silent synapses available for conversion by LTP to active synapses following SD and this availability is recovered by sleep.

The sleep-dependent electrophysiological findings raise several questions, including: 1) are other cell types besides the layer 2-3 pyramidal neurons from which we recorded, involved; 2) can changes in the transcriptome account, at least in part, for sleep-related synaptic changes in function, and if so, what are the specific functions of those genes and 3) what is the upstream sleep-dependent control of changes in gene expression?

Cell type/subtype specific gene expression in response to sleep loss

To examine the molecular mechanisms responsible for an SD response, we characterized single nuclei transcriptomes of cells from motor cortex of mice from two cohorts, an ad lib or control sleep (CS) group and an SD group (sleep deprived from ZT= 0-6hr; see detailed methods).

Removed in the revised version: Transcript libraries from each of the two cohorts (CS, n=4 and SD, n=4), were prepared from single nuclei isolated from mouse frontal cortex. Nuclei from each cohort were clustered and visualized using Uniform Manifold

Earlier version: Approximation and Projection (UMAP) and annotated using multi-modal characteristics from the Brain Initiative Cell Census(Network, 2021) (Fig. 2 A,B). The sleep condition had no significant effect on the expression-based clustering nor on the distribution and numbers of nuclei per cell sub-type across sleep conditions (Fig. 2 C). However, the SD condition was associated with a significantly increased number of transcripts/nucleus across sleep condition, matched for cell subtype (2way ANOVA, UMI/Nucleus(subType) x Condition(CSvsSD), F (16, 96) = 3.004, p=0.0004; Wilcoxon matched-pairs, p<= 0.0001; Table S2; Fig. 2 D). To avoid any technical biases, the CS and SD biological samples were treated in an identical manner, with sacrifice alternating between CS and SD samples, carried out by the same person, brain harvesting and frontal cortex isolation all carried out by Dr. Richa Pandey (blinded as to sample identity), RIN values were indistinguishable and finally, library preparation and sequencing were all completed in a single batch. Thus, the results may be indicative of a generalized increase in transcript number per cell sub-type.

Revised version: Transcript libraries from each of the two cohorts (CS, n=4 and SD, n=4), were prepared from single nuclei isolated from mouse frontal cortex. Nuclei from each cohort were clustered and visualized using Uniform Manifold Approximation and Projection (UMAP) and annotated using multi-modal characteristics from the Brain Initiative Cell Census(Network, 2021) (Fig. 2 A,B). The sleep condition had no significant effect on the expression-based clustering nor on the distribution and numbers of nuclei per cell sub-type across sleep conditions (Fig. 2 C). However, the SD condition was associated with a significantly increased number of transcripts/nucleus across sleep condition, matched for cell subtype (2way ANOVA, UMI/Nucleus(subType) x Condition(CSvsSD), F (16, 96) = 3.004, p=0.0004; Wilcoxon matched-pairs, p<= 0.0001; Table S2; Fig. 2 D). To avoid any technical biases, the CS and SD biological samples were treated in an identical manner, with sacrifice alternating between CS and SD samples, carried out by the same person, brain harvesting and frontal cortex isolation all carried out by Dr. Richa Pandey (blinded as to sample identity), RIN values were indistinguishable and finally, library preparation and sequencing were all completed in a single batch. Thus, the results may be indicative of a generalized increase in transcript number per cell sub-type.

Added in the revised version: snRNAseq data shows cell type and sub-type based on gene expression patterns are unaffected by sleep need.

Earlier version: snRNAseq data shows cell type and sub-type based on gene expression patterns are unaffected by sleep need. A) UMAP projection of cell-type gene expression pattern following 6 hours ad lib sleep (CS) at ZT=6hour. B) As in “A” except after 6hours sleep deprivation (SD), ZT=6hour. C) The distribution of cell numbers across subtypes is unaffected by sleep need. D) The median number of UMIs/cell is significantly increased by sleep need across all cell subtypes (see_Table S2 and text for statistics).

Revised version: A) UMAP projection of cell-type gene expression pattern following 6 hours ad lib sleep (CS) at ZT=6hour. B) As in “A” except after 6hours sleep deprivation (SD), ZT=6hour. C) The distribution of cell numbers across subtypes is unaffected by sleep need. D) The median number of UMIs/cell is significantly increased by sleep need across all cell subtypes (see_Table S2 and text for statistics).

We next determined the differential gene expression (DGE) of transcripts between conditions of CS and SD (see methods; Fig. 3 A). In response to 6 hours of SD from ZT= 0 to 6, both increased and decreased gene expression was observed in all cell subtypes (Table S2, S3). Predictably, there is a trend for increased DEGs in cell types expressing a greater number of genes (Fig. 3 A). The greatest proportion of sleep DEGs is ∼69%, expressed by ExIT cell types, while the next largest proportion is ∼17% for IN cell types (Fig. 3 B). To determine whether the increased numbers of genes expressed in the ExIT subtypes accounts for this cell type’s high proportion of DEGs, we examined each cell type’s probability of DEG expression given the number of genes expressed in the cell type, its Bayesian probability of expression.

Added in the revised version: Cell type specific differential gene expression in response to 6 hours of SD.

Earlier version: Cell type specific differential gene expression in response to 6 hours of SD. A) An XY bar plot of cell type specific DEGs (both up and down regulated) organized by cell type specific number of expressed genes (X axis) shows the greatest number of DEGs are found in excitatory pyramidal neurons that project within the telencephalon (ExIT). B) A pie chart of the distribution of DEGs amongst different cell types illustrates that the majority of sleep DEGs are expressed by ExIT neurons. C) An analysis of the cell type specific DEG occurrence shows the greatest probability of significant sleep loss gene response is found in ExIT cells by more than 3-fold compared to all other cell types.

Revised version: A) An XY bar plot of cell type specific DEGs (both up and down regulated) organized by cell type specific number of expressed genes (X axis) shows the greatest number of DEGs are found in excitatory pyramidal neurons that project within the telencephalon (ExIT). B) A pie chart of the distribution of DEGs amongst different cell types illustrates that the majority of sleep DEGs are expressed by ExIT neurons. C) An analysis of the cell type specific DEG occurrence shows the greatest probability of significant sleep loss gene response is found in ExIT cells by more than 3-fold compared to all other cell types.

The Bayesian probability of sleep loss DEGs, for cell subtype, is more than three-fold greater for excitatory intratelencephalically projecting neurons (ExIT subtype, pyramidal cells (Network, 2021)) compared to any other class of cells (Fig. 3 C). These observations suggest, with respect to transcriptomic mediated changes in the frontal cortex, the ExIT class neurons comprise the major target of sleep function and extends earlier observations that frontal glutamatergic neurons are the primary targets of the SD transcriptomic response (Bjorness et al., 2020; Kim et al., 2022).

We determined the functional relevance of cell-type specific DEGs with respect to both disease risk and biological properties. We found class-selective enrichment of both autism spectrum disorder genes (ASD; Fig. 4 A) and synaptic shaping components (Koopmans et al., 2019) (Fig. 4 A.).

Added in the revised version: DEG enrichment of cell types in response to SD by autism risk genes, synaptic shaping component genes and DEGs from MEF2c loss of function or constitutive HD4 repression of MEF2c.

Earlier version: DEG enrichment of cell types in response to SD by autism risk genes, synaptic shaping component genes and DEGs from MEF2c loss of function or constitutive HD4 repression of MEF2c. A) Heat map for cell type DEG enrichment by ASD risk genes, Synaptic Shaping Component genes, MEF2c-cKO DEGs and HD4cn DEGs. B) Model for the control of sleep DEGs by HD4 repression of MEF2c and by pMEF2c during low sleep need and, during high sleep need, de-repression of MEF2c by sequestration of pHD4 to the cytoplasm and dephosphorylation of MEF2c, facilitating expression of sleep genes. C) Model of downstream control by MEF2c and HD4 to facilitate SSC gene expression as sleep need increases. The SSC gene expression after waking, can then decrease AMPA/NMDA ratio and increase silent synapses during sleep to bias glutamate synapses towards decreased strength but increased potential for LTP at the start of the active phase when sleep need is low. D) An illustration of the SWS-SWA response to chronic MEF2c repression or activation. Chronic activation of MEF2c facilitated transcription leads to decreased AMPAR-mediated synaptic strength mimicking increased Ado tone that inhibits cortical-thalamic, glutamate synaptic activity, decreasing activation of cholinergic arousal centers and increasing SWS-SWA. Chronic repression of MEF2c does the opposite, mimicking loss of ADORA1 function and decreasing SWS-SWA.

Revised version: A) Heat map for cell type DEG enrichment by ASD risk genes, Synaptic Shaping Component genes, MEF2c- cKO DEGs and cnHD4 DEGs. B) Model for the control of sleep DEGs by HD4 repression of MEF2c and by pMEF2c during low sleep need (left). During high sleep need, MEF2c is de-repressed by sequestration of pHD4 to the cytoplasm and dephosphorylation of MEF2c. Both these events facilitate expression of sleep genes. C) It is proposed that SSC gene expression induced by prolonged waking or SD can, once asleep, decrease AMPA/NMDA ratio and increase silent synapses (SS) during sleep. This may bias glutamate synapses towards decreased strength but increased potential for LTP in preparation for the active phase, when sleep need is low. Conversely, as the active phase progresses, AMPA/NMDA ratio increases (as does synaptic strength), silent synapses are replaced by active synapses in association with increased expression of SSC genes to complete the cycle of glutamate, synapse, phenotype oscillation. D) An illustration of the SWS-SWA response to chronic MEF2c repression or activation. Chronic activation of MEF2c facilitated transcription leads to decreased AMPAR-mediated synaptic strength mimicking the effect of increased Ado tone, that will inhibit cortical-thalamic, glutamate synaptic activity to increase SWS-SWA. Chronic repression of MEF2c does the opposite, mimicking loss of ADORA1 function and tone, to decrease SWS-SWA.

Enrichment of sleep responsive DEGs by SSC and ASD genes, is strongest for ExIT (L 2-3, 4-5, 5) pyramidal neurons (Chi square with Yates correction p<0.006). In particular, a number of these sleep need ExIT-DEGs encode proteins that can directly affect AMPA/NMDA ratio and silent synapse numbers, including Dcc (Horn et al., 2013), Dgkb (Kakefuda et al., 2016), Gpc6 (Sato et al., 2016), Grin3a (Perez-Otano et al., 2016), Kif17 (Setou et al., 2000; Iwata et al., 2020), Kirrl3 (Martin et al., 2015) and Ptprf (Sclip and Sudhof, 2020). Some sleep modulated SSCs are also ASD risk factors in frontal cortical EXIT pyramidal neurons, including Cdh13, Dcc, Glra2, Gpc6, Grik4, Itpr1, Kirrel3, Nr1d1, Pcdh15 and Ush2a. To a much lesser extent, Layer 6 corticothalamic pyramidal neurons and somatostatin, parvalbumin and Sncg classes of inhibitory neurons (Scala et al., 2021) also show this enrichment of their sleep DEGs (Figure 4 A). These observations are consistent with the ExIT pyramidal cell, glutamatergic, synaptic phenotype as a major functional target for sleep proteins and implicates sleep (dis-)function’s role in autism risk.

Earlier version: The transcription factor, MEF2c is necessary for sleep loss gene expression (Bjorness et al., 2020). ExIT sleep loss genes are significantly (p<0.0001, Chi2 with Yates correction) enriched with MEF2c target genes (Harrington et al., 2016; Bjorness et al., 2020) (Figure 4 A.). Class II histone deacetylases 4/5 (HD4, HD5) are MEF2c binding partners and repressors (McKinsey et al., 2000), that can shuttle in and out of the nucleus and are trapped in the cytoplasm when phosphorylated by SIKinases (salt induced kinases (Kim et al., 2022; Zhou et al., 2022)), known to become active in conditions of high sleep need. As the mouse sleep phase progresses, HD4/5 is progressively de-phosphorylated in association with decreased sleep need and conversely as the mouse active phase progress, HD4/5 is progressively phosphorylated in association with increased sleep need (see Figure 4 in Zhou, et al. (Zhou et al., 2022)). A virally-mediated expression of mutant form of HD4, that resists phosphorylation and is thus at constitutively high levels in the nucleus (HD4cn (Zhou et al., 2022)), is expected to mimic a loss of function of MEF2c through constitutive repression of MEF2c. The sleep loss genes observed in frontal cortical ExIT neurons are enriched for both conditional Mef2c (Bjorness et al., 2020) loss of function DEGs (p<0.0001, Chi2 with Yates correction) and for DEGs in response to expression of HD4cn (Zhou et al., 2022) (p<0.0001, Chi2 with Yates correction), as illustrated in a sleep-transcriptome expression model (Figure 4 B.). Further, these findings are indicative of the role of HD4 as a repressor of MEF2c facilitated sleep gene expression when sleep need is low and its derepression of MEF2c through phosphorylation induced sequestration into the cytoplasm when sleep need is high (Zhou et al., 2022). The resulting high-sleep-need DEGs as noted above, can mediate the observed recovery of glutamate synapse phenotype from a high AMPA/NMDA ratio with sparse silent synapses, to lowered AMPA/NMDA ratio and increased silent synapses (Fig. 4 C).

Revised version: The transcription factor, MEF2c is necessary for sleep loss gene expression (Bjorness et al., 2020). ExIT sleep loss genes are significantly (p<0.0001, Chi2 with Yates correction) enriched with MEF2c target genes (Harrington et al., 2016; Bjorness et al., 2020) (Figure 4 A.). Class II histone deacetylases 4/5 (HD4, HD5) are MEF2c binding partners and repressors (McKinsey et al., 2000), that can shuttle in and out of the nucleus and are trapped in the cytoplasm when phosphorylated by SIKinases (salt induced kinases (Kim et al., 2022; Zhou et al., 2022)), known to become active in conditions of high sleep need. As the mouse sleep phase progresses, HD4/5 is progressively de-phosphorylated in association with decreased sleep need and conversely as the mouse active phase progress, HD4/5 is progressively phosphorylated in association with increased sleep need (see Figure 4 in Zhou, et al. (Zhou et al., 2022)). A virally-mediated expression of mutant form of HD4, that resists phosphorylation and is thus at constitutively high levels in the nucleus (HD4cn (Zhou et al., 2022)), is expected to mimic a loss of function of MEF2c through constitutive repression of MEF2c. The sleep loss genes observed in frontal cortical ExIT neurons are enriched for both conditional Mef2c (Bjorness et al., 2020) loss of function DEGs (p<0.0001, Chi2 with Yates correction) and for DEGs in response to expression of HD4cn (Zhou et al., 2022) (p<0.0001, Chi2 with Yates correction), as illustrated in a sleep-transcriptome expression model (Figure 4 B.). Further, these findings are indicative of the role of HD4 as a repressor of MEF2c facilitated sleep gene expression when sleep need is low and its de-repression of MEF2c through phosphorylation induced sequestration into the cytoplasm when sleep need is high (Zhou et al., 2022). The resulting high-sleep-need DEGs as noted above, can mediate the observed recovery of glutamate synapse phenotype from a high AMPA/NMDA ratio with sparse silent synapses, to lowered AMPA/NMDA ratio and increased silent synapses (Fig. 4 C).

Discussion

Added in the revised version: The major claims of this study are: 1) SD increases the AMPA/NMDA receptor ratio and RS restores it; 2) SD decreases silent synapses compared to CS and RS restores their number after SD; 3) the majority of SD-induced DEGs are found in ExIT cells (glutamate pyramidal neurons projecting within the telencephalon);4)ExIT SD-induced DEGs are enriched for genes encoding synaptic shaping components and for autism spectrum disorder risk and; 5) these DEGs are also enriched for DEGs induced by Mef2c loss of function restricted to forebrain glutamate neurons (ExIT cells comprise a subset of these) and by over-expression of constitutively nuclear HDAC4 that represses MEF2c transcriptional function. The last claim is consistent with an intracellular signaling model (presented as a hypothesis to be tested, in figure 4B).

Earlier version: Prolonged waking is associated with increased glutamatergic synaptic strength (Vyazovskiy et al., 2008; Liu et al., 2010; Bjorness et al., 2020) and increased number of functional AMPAR synapses (Bjorness et al., 2020). Both effects are reversed by recovery sleep (Bjorness et al., 2020). The evidence presented here suggests this wake/sleep homeostasis also reflects a waking induced increase in functional AMPA/NMDA ratio of glutamate synapses of the frontal cortex together with loss of silent glutamatergic synapses. The latter suggests a reduced availability of glutamate synapses for conversion from silent to active, reflecting saturation of LTP. Accordingly, it is reasonable to summarize the implications of sleep-related glutamate synaptic phenotype oscillation as: 1) prolonged waking induces a generalized synaptic strengthening coupled with a negative bias against potentiating plasticity as available slots for AMPARs saturate, and 2) a recovery by sleep (Fig. 4 C).

Revised version: Prolonged waking is associated with increased glutamatergic synaptic strength (Vyazovskiy et al., 2008; Liu et al., 2010; Bjorness et al., 2020) and increased number of functional AMPAR synapses (Bjorness et al., 2020). Both effects are reversed by recovery sleep (Bjorness et al., 2020). The evidence presented here suggests this wake/sleep homeostasis also reflects a waking induced increase in functional AMPA/NMDA ratio of glutamate synapses of the frontal cortex together with loss of silent glutamatergic synapses. The latter suggests a reduced availability of glutamate synapses for conversion from silent to active, reflecting saturation of LTP. Accordingly, it is reasonable to summarize the implications of sleep-related glutamate synaptic phenotype oscillation as: 1) prolonged waking induces a generalized synaptic strengthening coupled with a negative bias against potentiating plasticity as available slots for AMPARs saturate, and 2) a functional recovery by sleep (Fig. 4 C).

Added in the revised version: There is evidence that learning of a motor task or experience of forced altered motor activity can result in localized increases in NREM (slow wave sleep)-slow wave activity (Huber et al., 2004; Huber et al., 2006) in the motor cortex. Since SWS-SWA is considered a marker for sleep homeostasis, the altered motor activity induced increase of SWS-SWA was considered evidence for sleep-related function. Our earlier work has clearly shown that the treadmill method of SD increases frontal cortical SWS-SWA rebound, indicating a sleep-homeostatic process (Bjorness et al., 2016; Bjorness et al., 2020). Furthermore, we have also shown that this means of experimental SD causes similar glutamate synaptic changes as those observed using other means of SD like gentle handling (Liu, et al., JoNS 2010).

Learning of motor tasks that can be robustly and quantitatively monitored, like bird song in juvenile zebra finches (Deregnaucourt et al., 2005; Kollmorgen et al., 2020), may reflect the sleep/wake oscillation of bias for LTP and synaptic strengthening. At the start of the active phase, facilitated motor learning improves song performance to a level that appears to saturate by the end of the day. After a night’s sleep, performance is degraded but shows enhanced learning, facilitated by decreased AMPA/NMDA ratio and increased availability of synapses that can be readily potentiated. Since there are inevitable, day to day, unrelated variances in the environment that should not be learned, mitigation of over-fitting a learned task may require an oscillatory, stairstep-like, learning process. In rats, performance during learning of a novel motor task can display these similar characteristics, especially at the start of learning a novel task when exploration of a multi-dimensional learning space is at a premium (Kim et al., 2023), consistent with an underlying sleep/wake oscillation of glutamate synaptic phenotype.

Our unbiased examination of the ontology of ExIT DEGs indicates a response to sleep loss to modify ExIT glutamate synapses involves SSC DEGs. Our electrophysiological observations, now show more particularly, that the gene-expression modifications of ExIT neurons are coupled to SD-responses of, 1) increased AMPA/NMDA ratio; 2) decreased functional silent synapses and; 3) the previously described, increase in synaptic strength and functional number (Fig. 1 and 4).

The ontology of the ExIT SD-response transcriptome is also notable for enrichment by autism-related risk genes, many of which overlap with our curated category of SSCs. A similar overlap has been noted for social affiliative behaviors, synaptic adhesion molecules, other synaptic shaping molecules with autism/autism spectrum disorder (ASD) risk genes in the mammalian CNS (Taylor et al., 2020). The association of sleep disruption with autism is generally appreciated (Mazurek et al., 2019) but the focus was on loss of time spent asleep or disrupted sleep, rather than loss of sleep function. The overlapped SSC and ASD DEGs provide a molecular association between sleep’s functional role in motor cortex and ASD risk that captures ASD’s association with motor deficits observed in patients (Chukoskie et al., 2013) and in mouse autism models (Cording and Bateup, 2023).

Earlier version: The molecular mechanisms responsible for sleep/wake DEGs observed in ExIT neurons include two critical transcription factors, MEF2c (Bjorness et al., 2020) and HD4/5 (Kim et al., 2022; Zhou et al., 2022). All neuronal and non-neuronal frontal cortical transcriptomic changes are abolished by loss of function of Mef2c (Bjorness et al., 2020). The abolishment of the sleep transcriptome remained even after restriction of the Mef2c knockout to CamKII-expressing glutamatergic neurons of the forebrain. Importantly, MEF2c was necessary for both the SD-transcriptomic response and for SD-induced increase of glutamatergic synaptic strength and functional synaptic number as well as for ensuing sleep-mediated recovery (Bjorness et al., 2020). Thus, MEF2c has an essential role in facilitating the expression of sleep genes needed for functional recovery from loss of sleep.

Revised version: The molecular mechanisms responsible for sleep/wake DEGs observed in ExIT neurons include two critical transcription factors, MEF2c (Bjorness et al., 2020) and HD4/5 (Kim et al., 2022; Zhou et al., 2022). All neuronal and non-neuronal frontal cortical transcriptomic changes are abolished by loss of function of Mef2c. The abolishment of the sleep transcriptome remained even after restriction of the Mef2c knockout to CamKII-expressing glutamatergic neurons of the forebrain (Bjorness et al., 2020). Importantly, MEF2c was necessary for both the SD-transcriptomic response and for SD-induced increase of glutamatergic synaptic strength and functional synaptic number as well as for ensuing sleep-mediated recovery (Bjorness et al., 2020). Thus, MEF2c has an essential role in facilitating the expression of sleep genes needed for functional recovery from loss of sleep.

The activity of either of these factors to alter the sleep transcriptome, is controlled by their phosphorylation state. When sleep need is high, MEF2c is de-phosphorylated (Bjorness et al., 2020) (from pMEF2c to MEF2c) and Salt Induced Kinases (SIKs) phosphorylate HD4 and HD5 (pHD4/5), sequestering them to the cytoplasm (Kim et al., 2022; Zhou et al., 2022). Both the phosphorylation of HD4/5s (Miska et al., 1999) and de-phosphorylation of MEF2c (Zhu and Gulick, 2004), de-represses MEF2c transcriptional activity. The de-repression of MEF2c and its dephosphorylation are essential for SD-induced differential expression (Bjorness et al., 2020). The target genes of MEF2c activation can be inferred from the cortical, differential transcriptome induced by MEF2c loss of function (Harrington et al., 2016). These same genes overlap with the differentially expressed genes observed in response to overexpression of constitutively nuclear HD4 (HD4cn). This is predictable as HD4cn will constitutively bind and repress MEF2c in the nucleus, thus mimicking a MEF2c loss of function. Furthermore, in association with low sleep need, nuclear, and thus transcriptionally active (as a repressor), HD4, is more abundant in a de-phosphorylated state, relative to transcriptionally inactive, phosphorylated cytoplasmic pHD4 (Kim et al., 2022; Zhou et al., 2022). NMDAR activation negatively controls HD4 transcriptional activity and nuclear localization (Sando et al., 2012). Accordingly, as NMDAR activity accumulates during the active period, an expected accumulation of pHD4 in the cytoplasm occurs to derepress MEF2c activity. Thus, HD4 may act to repress MEF2c transcriptional facilitation of sleep genes, during low sleep need and derepress MEF2c as sleep need builds during the active period. We observed DEGs from both MEF2c loss of function and HD4cn also sleep DEGs consistent with their interaction driven by sleep need (Figure 4 C).

Both the role of MEF2c to mediate sleep gene expression and nuclear HD4 to repress MEF2c activity, lead to an apparent paradox. When HD4 repression of MEF2c is chronically reduced or lost, sleep need, as indicated by slow wave activity during slow wave sleep (SWS-SWA), is increased (Kim et al., 2022). If MEF2c promotes sleep gene expression, then sleep need should be reduced not increased by its de-repression. Conversely, when constitutively, nuclear HD4 is overexpressed, to chronically repress MEF2c, sleep need, as indicated by SWS-SWA, is unexpectedly, reduced (Zhou et al., 2022). But, what is actually being indicated by SWS-SWA?

Added in the revised version: SWS-SWA has long been employed as a marker for sleep need or intensity, primarily because of its strong correlation with previous time spent awake (Franken et al., 2001; Borbely et al., 2016). Despite this marker’s long history, the mechanisms responsible for its correlation to sleep need are not well understood.

Earlier version: SWS-SWA has long been employed as a marker for sleep need or intensity, primarily because of its strong correlation with previous time spent awake (Franken et al., 2001; Borbely et al., 2016). Despite this marker’s long history, the mechanisms responsible for its correlation to sleep need are not well understood. A likely candidate is adenosine (Ado) since, it is increased by sleep-need and is necessary for SWS-SWA rebound in response to SD (Porkka-Heiskanen et al., 1997; Bjorness et al., 2009; Bjorness et al., 2016). Ado is increased by glutamate synaptic activity acting on NMDARs in arousal centers. Ado activation of Ado A1 receptors (ADORA1) inhibits AMPAR synaptic activity (Brambilla et al., 2005; Greene et al., 2017). Adk, which encodes adenosine kinase in glial cells is the high affinity metabolizing enzyme of Ado. Glial knockout of Adk increases extracellular Ado and SWA in both wake and sleep (Bjorness et al., 2016). Rebound SWS-SWA in response to experimental SD increases Ado (Porkka-Heiskanen et al., 1997) and the rebound is blocked by a CamKII:Cre driven conditional knockout of Adora1 (Bjorness et al., 2016). Thus Ado tone, whether increased by prolonged waking activity or by glial knockdown of Adk, through its activation of ADORA1, inhibits glutamate synaptic activity, including in arousal centers, to promote SWS-SWA (Greene et al., 2017) (Figure 4 D).

Revised version: SWS-SWA correlates with either MEF2c and/or to ADORA1 activation. Either of these activations correlates with reduction of glutamate synaptic strength (illustrated in Figure 4 D). Ado activation of Ado A1 receptors (ADORA1) inhibits AMPAR synaptic activity (Brambilla et al., 2005; Greene et al., 2017). Adk, which encodes adenosine kinase in glial cells is the high affinity metabolizing enzyme of Ado. Glial knockout of Adk increases extracellular Ado and SWA in both wake and sleep (Bjorness et al., 2016). Rebound SWS-SWA in response to experimental SD increases Ado (Porkka-Heiskanen et al., 1997) and the rebound is blocked by a CamKII:Cre driven conditional knockout of Adora1 (Bjorness et al., 2016). Thus Ado tone, whether increased by prolonged waking activity or by glial knockdown of Adk, through its activation of ADORA1, inhibits glutamate synaptic activity, including in arousal centers, to promote SWS-SWA (Greene et al., 2017) (Figure 4 D).

Earlier version: Sleep, itself, induces a down-regulation of glutamate synaptic strength, synaptic numbers, and AMPA/NMDA ratio. Chronic facilitation of MEF2c-dependent sleep genes due to de-repression of MEF2c by HD4/5’s loss of function should induce a chronic down-regulation of glutamatergic synaptic strength and number (Bjorness et al., 2020; Kim et al., 2022; Zhou et al., 2022). The sleep-gene induced, chronic glutamate synaptic down-regulation can mimic the effect of a tonic increase of extracellular Ado to increase SWA as occurred following glial loss of function of Adk (Bjorness et al., 2016). Furthermore, synaptic glutamate activity in cholinergic arousal centers (Rainnie et al., 1994; Brambilla et al., 2005) can be reduced, decreasing cholinergic mediated thalamo-cortical activation, and increasing SWS-SWA.

Revised version: Sleep, itself, induces a down-regulation of glutamate synaptic strength, synaptic numbers, and AMPA/NMDA ratio. Chronic facilitation of MEF2c-dependent sleep genes due to de-repression of MEF2c by HDAC4/5’s loss of function should induce a chronic down-regulation of glutamatergic synaptic strength and number (Bjorness et al., 2020; Kim et al., 2022; Zhou et al., 2022). The sleep-gene induced, chronic glutamate synaptic down-regulation can mimic the effect of a tonic increase of extracellular Ado to increase SWA as occurred following glial loss of function of Adk (Bjorness et al., 2016). Furthermore, synaptic glutamate activity in cholinergic arousal centers (Rainnie et al., 1994; Brambilla et al., 2005) can be reduced, decreasing cholinergic mediated thalamo-cortical activation to additionally increase SWS-SWA.

Earlier version: Conversely, repression of MEF2c by de-phosphorylated HD4/5 (or by expression of phospho-dead HD4cn) reduces SWS-SWA (Zhou et al., 2022). In this case, glutamate synaptic activity is chronically increased (directly increasing C-T activation and indirectly driving cholinergic arousal center activity) by the direct effect of loss of sleep function, as evidenced by a loss of function of Mef2c in glutamate neurons that prevents sleep gene expression (Bjorness et al., 2020). This prevents the normal down-regulation of glutamate synaptic function (Bjorness et al., 2020) and can mimic the effect of a loss of Adora1 function on SWS-SWA (Bjorness et al., 2009). It predicts HD4cn’s repressive effect on sleep gene expression in response to SD, consistent with the enrichment of SD genes by both HD4cn and Mef2c-cKO DEGs, that we observed (Fig. 4 A).

Revised version: Conversely, repression of MEF2c by de-phosphorylated HD4/5 (or by expression of phospho-dead HD4cn) reduces SWS-SWA (Zhou et al., 2022). In this case, glutamate synaptic activity can be chronically increased by the direct effect of loss of sleep function to down-scale glutamate synaptic strength (Bjorness et al., 2020). This can mimic the effect of a loss of Adora1 function that reduces SWS-SWA (Bjorness et al., 2009). It also predicts HD4cn’s effect to alter sleep gene expression in response to SD, consistent with the enrichment of SD genes by both HD4cn and Mef2c-cKO DEGs, that we observed (Fig. 4 A).

Earlier version: Generally, facilitation or depression of SWA correlates with up or down scaling effects on cortical glutamate neurotransmission, respectively, even in the absence of a direct effects on sleep need (Figure 4D). Many reagents that reduce the excitability of glutamate pyramidal cells by various mechanisms, including anesthetics like isoflurane, barbiturates or benzodiazepines in addition to those activating ADORA1, increase SWA. Those agents that increase excitability of cortical glutamate pyramidal neurons, correlate with decreased SWA. Accordingly, SWA can be a marker for glutamate synaptic strength and number, modulated under normal conditions, by sleep function.

Revised version: Generally, facilitation or depression of SWA correlates with up or down scaling effects on cortical glutamate neurotransmission, respectively, even in the absence of a direct effect on sleep need (Figure 4D). Many reagents that reduce the excitability of glutamate pyramidal cells by various mechanisms, including anesthetics like isoflurane, barbiturates or benzodiazepines in addition to those activating ADORA1, increase SWA. Those agents that increase excitability of cortical glutamate pyramidal neurons, correlate with decreased SWA. Accordingly, SWA can be a marker for glutamate synaptic strength and number, modulated under normal conditions, by sleep function.

Earlier version: In summary, at the end of a long episode of waking, when sleep need is high, glutamate synapses of ExIT cells in the frontal cortex are electro-physiologically strengthened, increased in functional number, and show increased AMPA/NMDA ratio, together with decreased availability of silent synapses. The silent synapse conversion by LTP to active synapses is thus limited, creating a bias for increased strength at the expense of potentiating plasticity. In the ExIT neuronal nuclei, the transcription factor, MEF2c is de-repressed by dephosphorylation and sequestration of its co-repressors, the classII HDACs, pHD4/5, to the cytoplasm. This facilitates transcription of sleep genes, including those encoding SSCs controlling glutamatergic synaptic phenotype and ASD risk genes. Recovery sleep recovers the functional phenotype. These observations suggest a daily oscillation from lowered glutamatergic synaptic strength and increased bias for potentiating plasticity at the start of the active phase when sleep need is low, to increased synaptic strength and saturated plasticity at the end of the active phase, when sleep need is high. In conclusion, we have provided electrophysiological evidence for a wake/sleep oscillation of glutamate synapse phenotype that can mediate a glutamatergic strength/plasticity oscillating bias. We also show that this wake/sleep oscillation is likely to occur in glutamate synapses of ExIT pyramidal neurons and be mediated by a select set of synaptic shaping component genes, a significant number of which are also autism risk genes, whose expression is controlled by MEF2C and HD4 transcription factors. Finally, this study implicates a framework within which optimal cortical-dependent motor training can occur in a recursive incremental manner, facilitated by wake/sleep glutamate synapse phenotypical oscillation.

Revised version: In summary, at the end of a long episode of waking, when sleep need is high, glutamate synapses of ExIT cells in the frontal cortex are electro-physiologically strengthened, increased in functional number, and show increased AMPA/NMDA receptor ratio, together with decreased availability of silent synapses. The silent synapse conversion by LTP to active synapses is thus limited, creating a bias for increased strength at the expense of potentiating plasticity. In ExIT neuronal nuclei, the transcription factor, MEF2c is de-repressed by de-phosphorylation and sequestration of its co-repressors, the classII HDACs, pHDAC4/5, to the cytoplasm. This facilitates transcription of sleep genes, including those encoding SSCs controlling glutamatergic synaptic phenotype and ASD risk genes. Recovery sleep recovers the functional phenotype. These observations suggest a daily oscillation from lowered glutamatergic synaptic strength and increased bias for potentiating plasticity at the start of the active phase when sleep need is low, to increased synaptic strength and saturated plasticity at the end of the active phase, when sleep need is high. In conclusion, we have provided electrophysiological evidence for a wake/sleep oscillation of glutamate synapse phenotype that can mediate a glutamatergic strength/meta-plasticity oscillating bias. We also show that this wake/sleep oscillation is likely to occur in glutamate synapses of ExIT pyramidal neurons and be mediated by a select set of synaptic shaping component genes, a significant number of which are also autism risk genes, whose expression is controlled by MEF2C and HDAC4 transcription factors. Finally, this study implicates a framework within which optimal cortical-dependent motor training can occur in a recursive incremental manner, facilitated by wake/sleep glutamate synapse phenotypical oscillation.

Materials and Methods

Electrophysiology

All electrophysiological animal experiments were approved by the Animal Experimental Committee at the University of Tsukuba. Animals were housed on a 12:12-h light/dark cycle (with the light automatically turned on at 9:00) at a stable temperature (24.5 ± 0.5 °C), with free access to food and water. C57/Bl6 male mice aged 8 to 12 weeks were used.

Slice preparation and recording

Mice were fully anesthetized at ZT 6 with isoflurane, and their brains were promptly removed (<1 min) and placed in ice-cold artificial cerebrospinal fluid (ACSF), composed of 124 mM NaCl, 26 mM NaHCO3, 3mM KCl, 2 mM CaCl, 1 mM MgSO4, 1.25 mM KH2PO4, 10 mM glucose, 300 to 310 mOsm, equilibrated with 95% O2 and 5% CO2.

Brains were mounted on the stage of a vibrating microtome (Leica VT1200 S) with cyanoacrylate glue. Coronal sections (300 μm thick) were obtained and allowed to recover in ACSF at room temperature for at least 1 h. Slices were submerged in a perfusion chamber placed under an upright microscope (BX51WI; Olympus) fitted with a custom LED IR illumination and Nomarski interference contrast. Slices were superfused with ACSF, at a rate of 2 mL/min with Picrotoxin (100 micromolar) added to block GABAA receptor-mediated responses. Neurons in layers 2/3 of primary motor cortex with a pyramidal cell morphology were patched under visual control. The intracellular solution contained (in mM) 130 Cs gluconate, 1 mM EGTA, 10 mM Hepes, 4 mM MgATP, 0.3 mM NaGTP, and 5 mM NaCl. The pH was adjusted to 7.33 with CsOH and osmolarity to 285 to 300 mOsm. Borosilicate patch pipettes were pulled to an open-tip resistance of 2 to 4 MΩ. Signals were amplified and filtered at 10 kHz (Axopatch 700A; Molecular Devices) and then digitized at 20 kHz using customized routines in commercial software (IGORPro; WaveMetrics). Series resistance values for the recording pipette ranged between 8 and 15 MOhm and experiments with changes larger than 25% were not used for further analyses.

Evoked synaptic responses were produced through extracellular stimulation (0.1 ms) every 10 s with monopolar glass pipette electrodes filled with ACSF and placed in the vicinity (200-400 micrometers) lateral to the recorded neuron in layer 2/3. For the AMPA/NMDA amplitude ratio experiments stimuli amplitudes were 8.7 uA +/-3.5 uA (SD) and for the failure rate experiments, 1.4 uA +/-0.9 uA (SD).

AMPA/NMDA EPSC analyses

AMPA responses were recorded at a holding potential of -90 mV and mixed AMPA-NMDA responses at a holding potential of +50 mV. The -90 mV holding potential was chosen according to precedent (Myme et al., 2003). The hyperpolarized holding potential increases driving force and permits lower stimulus strength for the same response size – reducing the likelihood for polysynaptic responses. Experiments with multiple response peaks at -90 mV were not included in the analysis. The -90 mV holding potential also increases NMDA receptor Mg block - resulting in a minimally contaminated AMPA response.

The AMPA/NMDA ratio was determined according to an established protocol (Myme et al., 2003); briefly, at least 10 traces were averaged at -90 mV and +50 mV, the AMPA response was calculated as the peak inward current at -90 mV and the NMDA current was calculated at +50mV in a 5 ms window 40 ms after the AMPA peak.

Earlier version: Failure rates for AMPA and AMPA+NMDA EPSCs: Minimal stimulation was achieved by lowering the stimulation strength until clear successes and failures could be observed at -90 mV holding potential and stimulation strength was then kept constant for 5-10 min to allow the system to stabilize before failure ratios were determined. Relative numbers of silent and total (silent + active) synapse in an L2-3 pyramidal neuron evoked in a slice by minimal presynaptic stimulation are correlated with failure rates (FR=failures/number of stimuli). FR-90 = (1-prelease)As, and FR+50 = (1-prelease)As+Ss, where As=number active synapses (at -90mV) and Ss= number of silent synapses (only active at +50mV). Since prelease for each synapse is the same for -90mV and +50mV (i.e., assuming the same probability of failure at both potentials with each stimulus) then:

Revised version: Failure rates for AMPA and AMPA+NMDA EPSCs: Minimal stimulation was achieved by lowering the stimulation strength until clear successes and failures could be observed at -90 mV holding potential and stimulation strength was then kept constant for 5-10 min to allow the system to stabilize before failure ratios were determined. Relative numbers of silent and total (silent + active) synapse in an L2-3 pyramidal neuron evoked in a slice by minimal presynaptic stimulation are correlated with failure rates (FR= failures/number of stimuli). FR-90 = (1-prelease)As, and FR+50 = (1-prelease)As+Ss, where As=number active synapses (at -90mV) and Ss= number of silent synapses (only active at +50mV). Since prelease for each synapse is the same for -90mV and +50mV (i.e., assuming the same probability of failure at both potentials with each stimulus) then: Thus, Since (FR+50/FR-90)= (1-prelease)As+Ss / (1-prelease) As, as Ss-->0, (FR+50/FR-90)--> 1, the maximum ratio possible if prelease is constant. Our uncorrected data show (FR+50/FR-90)=∼1.3 in SD, which suggests that FR+50 is overestimated by at least a factor of 1.3 (likely, due to decreased input resistance at Vh=+50mV). Accordingly, we adjusted FR+50 by this factor = 1.3, for all sleep conditions.

Removed in the revised version: ln(FR+50) = ln (1-prelease)As+Ss

Removed in the revised version: = (As+Ss)*ln(1-prelease) ln(FR-90)= ln(1-prelease)As

Removed in the revised version: =As*ln(pfailure) Thus,

Removed in the revised version: ln(FR+50)/ln(FR-90) = (As+Ss)/As

Removed in the revised version: =1 + Ss/As.

Removed in the revised version: Since (FR+50/FR-90)= (1-prelease)As+Ss / (1-prelease)As, as Ss-->0, (FR+50/FR-90)--> 1, the maximum ratio possible if prelease is constant. Our uncorrected data show (FR+50/FR-90)=∼1.3 in SD, which suggests that FR+50 is overestimated by at least a factor of 1.3 (likely, due to decreased input resistance at Vh=+50mV). Accordingly, we adjusted FR+50 by this factor = 1.3, for all sleep conditions.

Sleep deprivation

Earlier version: All mice were accommodated for 1-2 weeks to a cage and treadmill, consisting of a bottomless plexiglass enclosure suspended over the treadmill with ad lib access to food and water. During this time, all animals were exposed to treadmill movement at ZT=2 and ZT=14 hours, for 30 minutes. For sleep deprivation (SD) the slow-moving treadmill (SMTM; 0.1KM/hour; see “video 1”) was engaged from ZT 0-6 hours and for recovery sleep (RS) from ZT 0-4. There was no treadmill engagement for RS from ZT 4-6 hours and for control sleep (CS), no treadmill engagement for the entire ZT 0-6 hours period. All animals were sacrificed at ZT=6 hours and tissue rapidly prepared for either electrophysiological (exVivo brain slices) or transcriptomic analysis as described below.

Revised version: All mice were accommodated for 1-2 weeks to a cage and treadmill, consisting of a bottomless plexiglass enclosure suspended over the treadmill with ad lib access to food and water. During this time, all animals were exposed to treadmill movement at ZT=2 and ZT=14 hours, for 30 minutes. For sleep deprivation (SD) the slow-moving treadmill (SMTM; 0.1KM/hour; see video) was engaged from ZT 0-6 hours and for recovery sleep (RS) from ZT 0-4. There was no treadmill engagement for RS from ZT 4-6 hours and for control sleep (CS), no treadmill engagement for the entire ZT 0-6 hours period. All animals were sacrificed at ZT=6 hours and tissue rapidly prepared for either electrophysiological (exVivo brain slices) or transcriptomic analysis as described below.

Single nucleus Transcriptome analysis

snRNA-seq library preparation

Nuclei for snRNA-seq were isolated from male mice, 8-10 weeks old. Briefly, the tissue was homogenized using a glass Dounce homogenizer in 2 ml of ice-cold lysis buffer (10 mM Tris-HCl, 10 mM NaCl, 3 mM MgCl2 and 0.1% Nonidet P40 Substitute) and was incubated on ice for 5 min. Nuclei were centrifuged at 500g for 5 min at 4 °C, washed with 4 ml ice-cold lysis buffer and incubated on ice for 5 min. Nuclei were centrifuged at 500g for 5 min at 4 °C. After centrifugation, the nuclei were resuspended in 500 μl of nucleus suspension buffer (NSB) containing 1× PBS, 1% BSA (no. AM2618, Thermo Fisher Scientific) and 0.2 U μl−1 RNAse inhibitor (no. AM2694, Thermo Fisher Scientific). The nucleus suspension was filtered through a 70-μm Flowmi cell strainer (no. H13680-0070, Bel-Art). Debris was removed with density gradient centrifugation using Nuclei PURE 2 M sucrose cushion solution and Nuclei PURE sucrose cushion buffer from the Nuclei PURE prep isolation kit (no. NUC201-1KT, Sigma Aldrich). Nuclei PURE 2 M sucrose cushion solution and Nuclei PURE sucrose cushion buffer were first mixed in a 9:1 ratio. A 500 μl volume of the resulting sucrose solution was added to a 2-ml Eppendorf tube. A 900 μl volume of the sucrose buffer was added to 500 μl of isolated nuclei in NSB. A 1,400 μl volume of nucleus suspension was layered to the top of the sucrose buffer. This gradient was centrifuged at 13,000g for 45 min at 4 °C. The pellet of nuclei was resuspended, washed once in NSB and filtered through a 70-μm Flowmi cell strainer (no. H13680-0070, Bel-Art). The concentration of nuclei was determined using 0.4% trypan blue (no. 15250061, Thermo Fisher Scientific), and was adjusted to a final concentration of 1,000 nuclei per microlitre with NSB.

Droplet-based snRNA-seq libraries were prepared using Chromium Single Cell 3′ v3.1 (1000121, 10x Genomics) according to the manufacturer’s protocol(Zheng et al., 2017). Libraries were sequenced using an Illumina NovaSeq 6000. snRNA-seq preprocessing and annotationRaw sequencing data was obtained as BCL files from the McDermott sequencing core at UT Southewestern. BCL files were demultiplexed using cellranger mkfastq (10X Genomics CellRanger suite v3.1.0). Resulting FASTQ files were then accessed for quality using FASTQC (v0.11.5). A reference mouse genome-annotation index was built for mouse genome (GRCm38p6) and Gencode annotation (vM17) using cellranger mkref (10X Genomics CellRanger suite v3.1.0). Quality passed FASTQ files were further aligned to reference mouse genome-annotation index and raw count matrices were generated using cellranger count (10X Genomics CellRanger suite v3.1.0). To remove ambient RNA contamination, we used CellBender’s remove-background (https://github.com/broadinstitute/CellBender) on the raw count matrix per sample. We note that without ambient RNA removal, glial cells were shown to be conspicuously contaminated with neuronal ambient RNAs(Caglayan et al., 2022). Also, potential doublets were discarded using DoubletFinder (https://github.com/chris-mcginnis-ucsf/DoubletFinder) tool. Ambient RNA and doublets cleaned data were used for downstream analysis. For each sample, nuclei with less than 20,000 UMIs and percentage of reads mapping to mitochondrial genes of less then 0.5 were retained. Individual samples per sleep condition were first integrated using Seurat’s (v3) integration approach (IntegrateData) and integrated datasets were further clustered (ScaleData, RunPCA, FindClusters) as described in Seurat’s (v3) integration vignette (https://satijalab.org/seurat/archive/v3.0/integration). Clustered data was then visualized (RunUMAP) using Uniform manifold approximation and projection embeddings (UMAP). Cluster-specific gene markers were identified (FindMarkers) and significant marker genes were enriched using Fisher exact test for cell-type markers defined in Brain Initiative Cell Census to annotate the cell-types. Cell-types were also confirmed by expression of canonical marker genes.

Differential gene expression tests were performed using edgeR-based pseudobulk approach across sleep conditions per cell-type. Significant differentially expressed genes (DEGs) were identified using absolute log2 fold change >= 0.1375 (10%) and a false discovery rate (FDR) <= 0.05. Further significant DEGs were enriched across defined gene classes (such as SSC genes, ASD genes, Mef2c-cKO genes and HD4cn genes) using R package SuperExactTest to functionally annotate the genes.

Statistics

Earlier version: The exact sample size (n) for each experimental group/condition, is provided as a discrete number and unit of measurement. Statistics were taken from distinct samples and not measured repeatedly, with the exception of the matched sampled evoked EPSCs recorded at Vh=-90mV and at Vh=+50mV from the same neuron as part of the analyses of AMPAR and NMDAR EPSC amplitudes and failure rates. Statistical details are provided in appropriate tables (Tables 1-3) in tabs labelled “Stats ….”, in the “Supplementary Information-Data/Tables”.

Revised version: The exact sample size (n) for each experimental group/condition, is provided as a discrete number and unit of measurement. Statistics were taken from distinct samples and not measured repeatedly, with the exception of the matched sampled evoked EPSCs recorded at Vh=-90mV and at Vh=+50mV from the same neuron as part of the analyses of AMPAR and NMDAR EPSC amplitudes and failure rates. Statistical details are provided in appropriate tables in tabs labelled “Stats ….”, in the “Supplementary Information-Data/Tables” (Tables 1, tabs 2,3,5,6; Table 2, tabs 2,3).